# SIDAS502: Math Methods for Data Science 


# School of Information, University of Michigan


## Week1: Vectors and Matrices Computation/Basic image processing 

### Version 1.1

This assignment is designed to let you apply basic concepts and perform basic algebra operations using vectors and matrices in Python. We will use Jupyter Notebooks to organize and submit your assignment. Please read the directions carefully, as we want to avoid submissions that are marked incorrect due to formatting mistakes. You will use vector and matrix computation using numpy and image processing with openCV. 

- Submit the following code at the beginning of your assignment:
- import pandas as pd
- import numpy as np
- from scipy.ndimage import zoom
- import unittest

Please enter your name: "Arthur Cho Ka Wai"

## Part 1: Vector  Computation

<strong>1.1</strong> \[1 pt\] Please write a function called <strong>computation</strong> that takes three variables called <strong>a(integer), v1(vector represented as a list)</strong> and <strong>v2(vector represented as a list)</strong> and returns a dictionary that has six keys: <strong>addition</strong> (the sum of v1 and v2, this should be a numpy array), <strong>subtraction</strong> (the difference from subtracting v2 from v1, this should be a numpy array), <strong>scale</strong> (scalling v1 by a, this should be a numpy array), <strong>dot</strong> (dot product of v1 and v2, this should be a float), <strong>length</strong> (norm of v1, this should be a float), and <strong>cross</strong> (cross product of v1 and v2, this should be a numpy array).

```python
import pandas as pd
import numpy as np
from scipy.ndimage import zoom
import unittest
```

<!-- grade_id: cell-22c5340398fc526a, points: 0, grade: false, solution: true -->
```python
def computation(a,v1,v2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f4a2e31983c60a0a, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 1.1 are within this cell
```

Below is the example of a=1, v1=[1,6] and v2=[8,-3] by the computation method. Please run the code, play around the numbers, and see the result.

```python
v1=[1,6]
v2=[8,-3]
result = computation(1,v1,v2)
print(result)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-f6ca6d631e70> in <module>
      1 v1=[1,6]
      2 v2=[8,-3]
----> 3 result = computation(1,v1,v2)
      4 print(result)

NameError: name 'computation' is not defined
```

<strong>1.2</strong> \[1 pt\] What is the dimension of v3? Please store the v3 dimension into variable <strong>v3_dimen</strong> using numpy.

For grading purposes, `v2_dimen` should be type `tuple`.

<!-- grade_id: cell-3c74940585fa0a97, points: 0, grade: false, solution: true -->
```python
v3 = np.array([5, 4])


# YOUR CODE HERE
#raise NotImplementedError()
v3_dimen
```

Output:
```
(2,)
```

<!-- grade_id: cell-667a0fabd0b24a44, points: 1, grade: true, solution: false -->
```python
assert type(v3_dimen) == tuple, "Problem 1.2, testing v3_dimen, type of value stored in variable does not match the expected type. Expecting Tuple."
```

<strong>1.3</strong> \[1 pt\] What is the v4 dimension? Please store the v4 dimension into variable <strong>v4_dimen</strong> using numpy.

For grading purposes, `v4_dimen` should be type `tuple`.

<!-- grade_id: cell-d41c577a699b60e7, points: 0, grade: false, solution: true -->
```python
v4 = np.array([1,  2,  3, 4, 5])

v4_dimen=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d32bdb8279f75fb8, points: 1, grade: true, solution: false -->
```python
assert type(v4_dimen) == tuple,  "Problem 1.3, testing v4_dimen, type of value stored in variable does not match the expected type. Expecting Tuple."
```

<strong>1.4</strong> \[1 pt\] Add the following two vectors by hand.  $v_1 = (5 * a, 4)$ and $v_2 (3*b, 7)$.  Store the results as strings in a python list called added (e.g. the vector $(5 * a, 4)$ would be ["5 * a", "4"].  Make sure that each entry in the list is a valid pyton expression.  So "5 * a" is correct, but "5 a" is not because "5 a" will not be properly evaluated by python.

<!-- grade_id: cell-02682c12e6797934, points: 0, grade: false, solution: true -->
```python
added = ['5 * a + 3 * b', '11']
# YOUR CODE HERE
#raise NotImplementedError()

```

<!-- grade_id: cell-f141102d74386afa, points: 1, grade: true, solution: false -->
```python
assert type(added) == list,  "Problem 1.4, testing added, type of value stored in variable does not match the expected type. Expecting list."
```

## Part 2: Matrix Computation 

<strong>2.1</strong> \[1 pt\] Write a function called <strong>Matrices</strong> that takes A matrix and B matrix, and then returns a dictionary that has four keys called <strong>addition</strong> (addition of A and B), <strong>subtraction</strong> (subtract B from A), <strong>multiplication</strong> (multiply A and B), <strong>elementwise</strong> (elementwise product, (a\*b)).  Inputs will be represented as numpy arrays and outputs should also be numpy arrays.     

<!-- grade_id: cell-f34eaaa22752fde4, points: 0, grade: false, solution: true -->
```python
def Matrices(A,B):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e8861c2001b4e0ff, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.1 are within this cell
```

<strong>2.2</strong> \[1 pt\] Assume A=[[-2,6,8],[5,-2,8],[2,1,4]] and B=[[10,-4,6],[1,8,-3],[8,-2,4]]. What are the addition, subtraction, multiplication, and elementwise product of A matrix and B matrix? Please store your answers into a variable <strong>ANS22</strong> formatted as a dictionary as in problem 2.1.  (Use the function from 2.1) 


<!-- grade_id: cell-d59d72a05083a65b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1943466bf97bdfa1, points: 1, grade: true, solution: false -->
```python
assert type(ANS22) == type({}), "Problem 2.2, testing ANS22, type of value stored in variable does not match the expected type. Expecting Dictionary."
```

<strong>2.3</strong> \[1 pt\] What is the transpose of C? Please store the transpose of C into a variable <strong>C_trans</strong>. 

For grading purposes, `C_trans` should be a numpy array.

<!-- grade_id: cell-0afc4735301d7086, points: 0, grade: false, solution: true -->
```python
C = np.array([[1, 1, 2], [3, 5, 8]])

C_trans=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2ae88de7677e6d09, points: 1, grade: true, solution: false -->
```python
assert type(C_trans) == np.ndarray, "Problem 2.3, testing C_trans, type of value stored in variable does not match the expected type. Expecting np.ndarray"
```

<strong>2.4</strong> \[1 pt\] What is (CD)<sup>T</sup>? Please store the value into <strong>CD_trans.</strong>

For grading purposes, `CD_trans` should be type `numpy.array`.

<!-- grade_id: cell-71cbc2342eada7dc, points: 0, grade: false, solution: true -->
```python
D = np.array([[13], [21], [34]])


CD_trans=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-174e47222abd77ec, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.4 are within this cell
assert type(CD_trans) == np.ndarray, "Problem 2.4, testing CD_trans, type of value stored in variable does not match the expected type. Expecting np.ndarray"
```

<strong>2.5</strong> \[1 pt\] What is the inverse matrix of E? Please store it into a variable <strong>E_inv</strong>.              

For grading purposes, `E_inv` should be type `numpy.matrix`

<!-- grade_id: cell-98a08d219887a362, points: 0, grade: false, solution: true -->
```python
E = np.array([[1, 1, 2], [3, 5, 8], [13, 21, 50]])

E_inv=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-75c20b117fcb21c5, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.5 are within this cell
```

<strong>2.6</strong> \[1 pt\] If AD=I, what is matrix D? Please store matrix D into a variable <strong>D</strong>.  (Use A from problem 2.2).

For grading purposes, `D` should be type `numpy.ndarray`.

<!-- grade_id: cell-f9ac76d940b23358, points: 0, grade: false, solution: true -->
```python
D= None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7e7ef1859427aebc, points: 1, grade: true, solution: false -->
```python
assert type(D) == np.ndarray,  "Problem 2.6, testing D, type of value stored in variable does not match the expected type. Expecting np.array"
```

<strong>2.7</strong> \[1 pt\] What is the determinant of A? Please store the value into a variable <strong>A_deter</strong>.  (Use A from problem 2.2.) 

For grading purposes, `A_deter` should be type `numpy.float64`.

<!-- grade_id: cell-ded3892aaac45733, points: 0, grade: false, solution: true -->
```python
A = np.array([[-2,6,8],[5,-2,8],[2,1,4]])

A_deter=None 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0a36f464f9700b09, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.7 are within this cell
```

<strong>2.8</strong> \[1 pt\] What is the determinant of BA? Please store the value into a variable <strong>BA_deter</strong>.  (Use A and B from problem 2.2.)

For grading purposes, `BA_deter` should be type `numpy.float64`.

<!-- grade_id: cell-f41d8e11e74b3fc0, points: 0, grade: false, solution: true -->
```python

BA_deter=None 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4a21c6f652c4f3dd, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.8 are within this cell
```

<strong>2.9</strong> \[1 pt\] Change negative values of A to 1 and store the result into a variable <strong>A_pos</strong>.   (Use A from problem 2.2.)  Hint:boolean matrix 

For grading purposes, `A` should be type `numpy.matrix`.

<!-- grade_id: cell-45fbe97d7f713f8a, points: 0, grade: false, solution: true -->
```python
A_pos = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-aa0cf2d089271312, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.9 are within this cell
```

<strong>2.10</strong> \[1 pt\] We are going to find the trace of the sum of two scaled matrices. Find the trace of 2\*A + 3\*B and store the value into a variable <strong>trace2A3B</strong>.  (Use A and B from problem 2.2.)

For grading purposes, `trace2A3B` should be type `numpy.int64`.

<!-- grade_id: cell-4a58204752a1e3a5, points: 0, grade: false, solution: true -->
```python


trace2A3B=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-09f9e02c3fb55c0d, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.10 are within this cell
```

<strong>2.11</strong> \[1 pt\] What is the value of tr(AB<sup>T </sup>)? Please store the value into a variable <strong>AB_trace</strong>.  (Use A and B from problem 2.2).

For grading purposes, `AB_trace` should be type `numpy.int64`.

<!-- grade_id: cell-03d1b2f808aa170a, points: 0, grade: false, solution: true -->
```python
AB_trace = None 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-12f5c58920c36677, points: 1, grade: true, solution: false -->
```python
assert type(AB_trace) == np.int64, "Problem 2.11, testing AB_trace, type of value stored in variable does not match the expected type. Expecting np.int64"
```

## Part 3: Matrix Calculations 

<strong>3.1</strong> \[1 pt\] Let $F = [[1, 2, a], [3, 2, 1 + b], [2, 2, 2 + c]]$ where $a$, $b$, and $c$ are variables.  $G = [[0, 2 + c, -a], [3, b + c, -1], [a, 3, -b]]$ where $a$, $b$, and $c$ are the same variables as in $F$.   What is the value of $F + G$? Please store the value into a string <strong>FG_sum</strong> written with valid python code formatting (e.g. FG_sum = "[[1, 2, a], [3, 2, 1 + b], [2, 2, 2 + c]]").  (Note you are encouraged to do this by hand.)

<!-- grade_id: cell-69904b9325be7f33, points: 0, grade: false, solution: true -->
```python
FG_sum = None 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0b6fa79caf358a18, points: 1, grade: true, solution: false -->
```python
assert type(FG_sum) == str, "Problem 3.1, type of value stored in variable FG_summ does not match the expected type."
```

<strong>3.2</strong> \[1 pt\] Assume F and G as above.  What is the (matrix) product of F and G? Please store your answers into a string variable <strong>Prod_FG</strong>  written with valid python code formatting (e.g. FG_sum = "[[1, 2, a], [3, 2, 1 + b], [2, 2, 2 + c]]"). (Note you are encouraged to do this by hand.)


<!-- grade_id: cell-d59d72a05083a65a, points: 0, grade: false, solution: true -->
```python
Prod_FG = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1943466bf97bdfa2, points: 1, grade: true, solution: false -->
```python
assert type(Prod_FG) == str, "Problem 3.2, type of value stored in variable Prod_FG does not match the expected type."
```

<strong>3.3</strong> \[1 pt\] What is the transpose of F?  Please store your answers into a string variable <strong>F_trans</strong>  written with valid python code formatting (e.g. F_trans = "[[1, 2, a], [3, 2, 1 + b], [2, 2, 2 + c]]"). (Note you are encouraged to do this by hand.)

<!-- grade_id: cell-0afc4735301d7087, points: 0, grade: false, solution: true -->
```python
F_trans=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2ae88de7677e6d00, points: 1, grade: true, solution: false -->
```python
assert type(F_trans) == str, "Problem 3.3, type of value stored in variable F_trans does not match the expected type."
```

 <strong>3.4</strong> \[1 pt\] What is the determinant of F? Please store your answers into a string variable <strong>Det_F</strong>  written with valid python code formatting (e.g. F_trans = "2 + a -b"). (Note you are encouraged to do this by hand.)

<!-- grade_id: cell-71cbc2342eada9dc, points: 0, grade: false, solution: true -->
```python
Det_F=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-174e47222abd77e7, points: 1, grade: true, solution: false -->
```python
assert type(Det_F) == str, "Problem 3.4, type of value stored in variable Det_F does not match the expected type."
```

<strong>3.5</strong> \[1 pt\] What is the trace of the matrix 3F + 2G?    Please store your answers into a string variable <strong>Trace_FG</strong>  written with valid python code formatting (e.g. Trace_FG = "2 + a -b"). (Note you are encouraged to do this by hand.)


<!-- grade_id: cell-98a08d219887a363, points: 0, grade: false, solution: true -->
```python
import numpy as np

Trace_FG=None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-75c20b117fcb21c3, points: 1, grade: true, solution: false -->
```python
assert type(Trace_FG) == str, "Problem 3.5, type of value stored in variable Trace_FG does not match the expected type."
```

<strong>3.6</strong> \[1 pt\] What is the trace of the matrix $F^T \cdot  G^T$?    Please store your answers into a string variable <strong>Trace_FTGT</strong>  written with valid python code formatting (e.g. Trace_FTGT = "2 + a -b"). (Note you are encouraged to do this by hand.)

Hint: no/very little additional calculation is necessary! 


<!-- grade_id: cell-f9ac76d940b23353, points: 0, grade: false, solution: true -->
```python
Trace_FTGT= None
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-7e7ef1859427aeb3, points: 1, grade: true, solution: false -->
```python
assert type(Trace_FTGT) == str, "Problem 3.6, type of value stored in variable Trace_FTGT does not match the expected type."
```

# Part 4: Image 

OpenCV is an open source computer vision library(https://opencv-python-tutroals.readthedocs.io/en/latest/py_tutorials/py_setup/py_intro/py_intro.html#intro). It is used in deep learning or image processing. Here, we will import cv2 and then use it to read the Michigan_logo and plot the image. If you are unfamilair with the library, you may need to read some of the documentation.




```python
import cv2
from matplotlib import pyplot as plt
```

<strong>4.1</strong> Please use pyplot to read the image and plot the image below. 

<!-- grade_id: cell-0a09e5f377a407e3, points: 0, grade: true, solution: true -->
```python
filesource = 'assets/Michigan_logo.jpg'
img=None 
# YOUR CODE HERE
raise NotImplementedError()
```

<strong>4.2</strong> \[1 pt\] What are the RGB values of the pixel at 10, 15 within the image? Please store your answer in a string format: <strong>ANS42</strong> = "R:, G:, B:". (example, string may look like the following “R:13, G:2, B:100”.) (Hint: use the python format function)

<!-- grade_id: cell-d1f5f261ccc5d800, points: 0, grade: false, solution: true -->
```python
ANS42= ""
# YOUR CODE HERE
raise NotImplementedError()
print(ANS42)
```

<!-- grade_id: cell-d2d5230f9865ca70, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 4.2 are within this cell
```

<strong>4.3</strong> Please resize the image to 20x20 using cv2.

<!-- grade_id: cell-8b16589ebbe0702a, points: 0, grade: true, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()
```

<strong>4.4</strong> \[1 pt\] What is the dimension of the image? Please store the result into the variable <strong>resized_dim</strong>. 

<!-- grade_id: cell-da73d6771390566c, points: 0, grade: false, solution: true -->
```python
resized_dim = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-214cff6388d2e1c8, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 4.4 are within this cell
```